# BT01: Ẩn tin mật trên văn bản bằng phương pháp chèn khoảng trắng <font color="grey">(chèn khoảng trắng giữa các từ để làm cho văn bản được căn lề 2 bên)


Họ tên: Nguyễn Đỗ Chí Thảo

MSSV: 1712159

---
## Cách làm bài và nộp bài
&#9889; Bạn lưu ý là mình sẽ dùng chương trình hỗ trợ chấm bài nên bạn cần phải tuân thủ chính xác qui định mà mình đặt ra, nếu không rõ thì hỏi, chứ không nên tự tiện làm theo ý của cá nhân.

**Cách làm bài**

Bạn sẽ làm trực tiếp trên file notebook này. Đầu tiên, bạn điền họ tên và MSSV vào phần đầu file ở bên trên. Trong file, bạn làm bài ở những chỗ có ghi là:
```python
# YOUR CODE HERE
raise NotImplementedError()
```
hoặc đối với những phần code không bắt buộc thì là:
```python
# YOUR CODE HERE (OPTION)
```
hoặc đối với markdown cell thì là:
```markdown
YOUR ANSWER HERE
```
Tất nhiên, khi làm thì bạn xóa dòng `raise NotImplementedError()` đi.
Đối những phần yêu cầu code thì thường ở ngay phía dưới sẽ có một (hoặc một số) cell chứa các bộ test để giúp bạn biết đã code đúng hay chưa; nếu chạy cell này không có lỗi gì thì có nghĩa là qua được các bộ test. Trong một số trường hợp, các bộ test có thể sẽ không đầy đủ; nghĩa là, nếu không qua được test thì là code sai, nhưng nếu qua được test thì chưa chắc đã đúng hoàn toàn.

Trong khi làm bài, bạn có thể cho in ra màn hình, tạo thêm các cell để test. Nhưng khi nộp bài thì bạn xóa các cell mà bạn tự tạo, xóa hoặc comment các câu lệnh in ra màn hình. Bạn lưu ý <font color=red>không được tự tiện xóa các cell hay sửa code của Thầy</font> (trừ những chỗ được phép sửa như đã nói ở trên).

Trong khi làm bài, thường xuyên `Ctrl + S` để lưu lại bài làm của bạn, tránh mất mát thông tin.


*Nên nhớ mục tiêu chính ở đây là <font color=green>học, học một cách chân thật</font>. Bạn có thể thảo luận ý tưởng với bạn khác, nhưng <font color=green>code và bài làm phải là của bạn, dựa trên sự hiểu thật sự của bạn</font>. <font color=red>Nếu vi phạm thì sẽ bị 0 điểm cho toàn bộ môn học.</font>*

**Cách nộp bài**

Khi chấm bài, đầu tiên mình sẽ chọn `Kernel` - `Restart Kernel & Run All Cells`, để restart và chạy tất cả các cell trong notebook của bạn; do đó, trước khi nộp bài, bạn nên chạy thử `Kernel` - `Restart Kernel & Run All Cells` để đảm bảo mọi chuyện diễn ra đúng như mong đợi.

Sau đó, bạn tạo thư mục nộp bài theo cấu trúc sau:
- Thư mục `MSSV` (vd, nếu bạn có MSSV là 1234567 thì bạn đặt tên thư mục là `1234567`)
    - Thư mục `BT01`
        - File `BT01-AnTinMatTrenVanBan_KhoangTrang.ipynb` (không cần nộp các file khác)

Cuối cùng, bạn nén thư mục `MSSV` này lại và nộp ở link trên moodle. <font color=red>Bạn lưu ý tuân thủ chính xác cấu trúc này.</font>

## 2. Import

```python
from bitarray import bitarray
```

## 3. Hàm nhúng (6đ)

<!-- grade_id: cell-2f21b430087c5c35, points: 0, grade: false, solution: true -->
```python
def embed(msg_file, cover_text_file, text_width, stego_text_file):
    '''
    Nhúng tin mật vào văn bản bằng phương pháp chèn khoảng trắng 
    để văn bản được căn lề 2 bên.
    
    Các tham số:
        msg_file (str): Tên file chứa secret message.
        cover_text_file (str): Tên file chứa cover text.
        text_width (int): Chiều dài của dòng sau khi được căn lề.
        stego_text_file (str): Tên file chứa stego text.
    Giá trị trả về:
        bool: True nếu nhúng thành không, False nếu không đủ chỗ để nhúng. 
    '''
    # YOUR CODE HERE
    
    # Read cover text file
    with open(cover_text_file, 'r') as f:
        cover_text_lines = f.read().splitlines()
    
    # Read msg file
    with open(msg_file, 'r') as f:
        msg = f.read()
        
    # Convert msg to msg bits
    msg_bits = bitarray()
    msg_bits.frombytes(msg.encode('utf-8'))    
    # Embed msg bits into cover text lines
    stego_text = ''; b = 0
    flag=False
    interation=0
    for i in range(len(cover_text_lines)-1):
        line=cover_text_lines[i]
        next_line=cover_text_lines[i+1]
        #--------------------
        word=line.split(' ')
        #so khoang trang giua cac tu
        n=len(word)-1
        #TH dòng trống || cuối đoạn||dòng có 1 từ
        if line==''or next_line=='' or n<=0:
            stego_text += line+'\n'
            continue
         #--------------------
        else:
            #so khoang trang can chen
            m = text_width - len(line) 
            # TH co the chen duoc bit trong msg_bits
            if 0<m and m< n:
                #So bit chen duoc
                n_bits=min(m,n-m)
                # Duyet n_bits
                for i in range(2*n_bits)[::2]:
                    #Chua duyet het bit trong msg_bits
                    if b < len(msg_bits):
                        if msg_bits[b]==1:
                            stego_text+=word[i]+' '+word[i+1]+'  '
                        else:
                            stego_text+=word[i]+'  '+word[i+1]+' '
                        b+=1
                    # Da duyet het msg_bits
                    else:
                        if not(flag):
                            flag=True
                            stego_text+=word[i]+' '+word[i+1]+'  '
                        else:
                            stego_text+=word[i]+'  '+word[i+1]+' '
                # Can le nhung khoang trang con lai khi k the chen them bit
                for i in range(2*n_bits,len(word)-1):
                    stego_text+=word[i]+' '
                    if n_bits<m:
                        stego_text+=' '
                        n_bits+=1
                stego_text+=word[-1]
            #TH khong the chen bit nhung van phai can le
            else:
                #----------------------------
                n_spaces=int(m/n)+1
                odd_spaces=int(m%n)
                for i in range(len(word)-1):
                    stego_text+=word[i]+' '*n_spaces
                    if odd_spaces>0:
                        stego_text+=' '
                        odd_spaces-=1
                stego_text+=word[-1]
                #-----------------------------
        stego_text +='\n'
        interation+=1
    if not(flag):
        return False
    # Write stego text to file
    stego_text+=cover_text_lines[len(cover_text_lines)-1]
    with open(stego_text_file, 'w') as f:
        f.write(stego_text)
    return True
```

<!-- grade_id: cell-057a58bf038e1e60, points: 5, grade: true, solution: false -->
```python
# TEST 1            
result = embed('msg.txt', 'cover.txt', 70, 'stego.txt')
assert result == True
with open('stego.txt', 'r') as f:
    stego_text = f.read()
with open('correct_stego.txt', 'r') as f:
    correct_stego_text = f.read()
assert stego_text == correct_stego_text
```

<!-- grade_id: cell-f625774497df6542, points: 1, grade: true, solution: false -->
```python
# TEST 2
result = embed('msg2.txt', 'cover.txt', 70, 'stego.txt')
assert result == False
```

## 4. Hàm rút trích (4đ)

<!-- grade_id: cell-f3f9b34770f282aa, points: 0, grade: false, solution: true -->
```python
def extract(stego_text_file,extr_msg_file):
    '''
    Hàm rút trích tin mật đã được nhúng bằng phương pháp chèn khoảng trắng.
    
    Các tham số:
        stego_text_file (str): Tên file chứa stego text.
        extr_msg_file (str): Tên file chứa secret message được rút trích.
    '''
    # YOUR CODE HERE
    with open(stego_text_file, 'r') as f:
        cover_text_lines = f.read().splitlines()
        
    mesg=bitarray()
    for test in cover_text_lines:
        #Lay vi tri cac khoang trang trong 1 dong
        l=[]
        for i, c in enumerate(test):
            if c==' ':
                l.append(i)
                
        #Convert tu vi tri sang list cac space tuong ung
        string=[]
        stg=' '
        for i in range(len(l)-1):
            if l[i+1]==l[i]+1:
                stg+=' '
            else:
                string.append(stg)
                stg=' '
        string.append(stg)
        
        #Convert from space array to bit array
        if len(string)>=2 and (\
                               (len(string[0])==2 and len(string[1])==1) or (len(string[0])==1 and len(string[1])==2)):
            for i in range(len(string)-1)[::2]:
                if len(string[i])==2 and len(string[i+1])==1:
                    mesg.append(False)
                elif len(string[i])==1 and len(string[i+1])==2:
                    mesg.append(True)
                else:
                    break
    for i in range(len(mesg)-1, -1, -1):
        bit=mesg.pop()
        if bit==True:
            break
    mesg=mesg.tobytes()
    extr_mesg=mesg.decode("utf-8") 
    with open(extr_msg_file, 'w') as f:
        f.write(extr_mesg)
```

<!-- grade_id: cell-caf1afe26d95b670, points: 4, grade: true, solution: false -->
```python
# TEST
extract('correct_stego.txt', 'extr_msg.txt')
with open('extr_msg.txt', 'r') as f:
    extr_msg = f.read()
with open('correct_extr_msg.txt', 'r') as f:
    correct_extr_msg = f.read()
assert extr_msg == correct_extr_msg
```